---
title: 'Basics of Hierarchical Queries'
description: ""
---

## Introduction to hierarchical data

```yaml
type: VideoExercise
key: 7f3c09d6be
xp: 50
```

`@projector_key`
3bffc663d01c6a7472521595d82a4c6a

---

## Working with CTEs  and hierarchical queries

```yaml
type: VideoExercise
key: b60d7a108e
xp: 50
```

`@projector_key`
6a229a22cf861b184050cf100cd889dc

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 5201537197
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
